## Module 4

In this assignment, you will implement an online banking system.  Users can sign-up with the system, log in to the
system, change their password, and delete their account.  They can also update their bank account balance and transfer
money to another user’s bank account.

You’ll implement functions related to File I/O and dictionaries.  Two of the functions require you to import files
and create dictionaries.  User information will be imported from the “users.txt” file and account information will be
imported from the “bank.txt” file.  Take a look at the content in the different files.  The remaining functions require
you to use or modify the two dictionaries created from the files.

Each function has been defined for you, but without the code. See the docstring in each function for instructions on
what the function is supposed to do and how to write the code. It should be clear enough.  In some cases, we have
provided hints to help you get started.


<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import nose.tools as tools
```

<!-- grade_id: import_and_create_bank, points: 0, grade: false, solution: true -->
```python
def import_and_create_bank(filename):
    # create an empty bank dictionary
    bank = {}
    # read the file
    f = open(filename, "r")
    # since every line represents one user, apply readlines() to create a list
    lines = f.readlines()
    # iterate over all the lines
    for line in lines:
        # strip white spaces and split line into list based on ":"
        lst = line.strip().split(":")
        # skip the line with only key or value
        if len(lst) <= 1:
            continue
        # assign name to key and deposit to value
        key = lst[0].strip()
        value = lst[1].strip()
        # skip if the deposit is not number
        try:
            value = float(value)
            bank[key] = bank.get(key, 0) + value
        except:
            continue
    # print(bank)
    # close the file
    f.close()
    return bank
```

<!-- grade_id: test_import_and_create_bank, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")

tools.assert_false(len(bank) == 0)
tools.assert_almost_equal(115.5, bank.get("Brandon"))
tools.assert_almost_equal(128.87, bank.get("James"))
tools.assert_is_none(bank.get("Joel"))
tools.assert_is_none(bank.get("Luke"))
tools.assert_almost_equal(bank.get("Sarah"), 827.43)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: signup, points: 0, grade: false, solution: true -->
```python
def check_upper(password):
    for char in password:
        if char.isupper():
            return True
    return False

def check_lower(password):
    for char in password:
        if char.islower():
            return True
    return False

def check_number(password):
    for char in password:
        if char.isdigit():
            return True
    return False

def validation(user_accounts, username, password):
    """test whether the username and the password are both valid.
    User_accounts is a dictionary contains all user names"""
    if (username not in user_accounts) and len(password) >= 8 and check_upper(password) and check_lower(password) and check_number(password) and username != password:
        return True
    else:
        return False


def signup(user_accounts, log_in, username, password):

    if validation(user_accounts, username, password):
        user_accounts[username] = password
        log_in[username] = False
        return True
    else:
        return False

```

<!-- grade_id: import_and_create_accounts, points: 0, grade: false, solution: true -->
```python
def import_and_create_accounts(filename):
    # Create an empty user accounts dictionary and an empty login dictionary
    user_accounts = {}
    log_in = {}
    # read the file
    f = open(filename, "r")
    # since every line represents one user, apply readlines() to create a list
    lines = f.readlines()
    # iterate over all the lines
    for line in lines:
        # strip white spaces and split line into list based on ":"
        lst = line.strip().split("-")
        # skip the line with only key or value
        if len(lst) <= 1:
            continue
        # assign name to key and deposit to value
        key = lst[0].strip()
        value = lst[1].strip()
        if validation(user_accounts, key, value):
            user_accounts[key] = value
            log_in[key] = False
        else:
            continue
    f.close()
    #print(user_accounts)
    #print(log_in)
    return user_accounts, log_in

```

<!-- grade_id: test_import_and_create_accounts, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(len(user_accounts) == 0)
tools.assert_false(len(log_in) == 0)
tools.assert_equal(user_accounts.get("Brandon"),"brandon123ABC")
tools.assert_equal(user_accounts.get("Jack"),"jack123POU")
tools.assert_is_none(user_accounts.get("Jennie"))
tools.assert_false(log_in["Sarah"])
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: test_signup, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(signup(user_accounts,log_in,"Brandon","123abcABCD"))

tools.assert_false(signup(user_accounts,log_in,"BrandonK","123ABCD"))
tools.assert_false(signup(user_accounts,log_in,"BrandonK","1234ABCD"))
tools.assert_false(signup(user_accounts,log_in,"BrandonK","abcdABCD"))
tools.assert_false(signup(user_accounts,log_in,"BrandonK","1234abcd"))

tools.assert_false(signup(user_accounts,log_in,"123abcABCD","123abcABCD"))

tools.assert_true(signup(user_accounts,log_in,"BrandonK","123aABCD"))
tools.assert_false(signup(user_accounts,log_in,"BrandonK","123aABCD"))
tools.assert_true("BrandonK" in user_accounts)
tools.assert_equal("123aABCD",user_accounts["BrandonK"])
tools.assert_false(log_in["BrandonK"])
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: login, points: 0, grade: false, solution: true -->
```python
def login(user_accounts, log_in, username, password):
    if (username not in user_accounts) or (password != user_accounts[username]):
        return False
    else:
        log_in[username] = True
        return True
```

<!-- grade_id: test_login, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(login(user_accounts, log_in,"Brandon","123abcAB"))
tools.assert_true(login(user_accounts, log_in,"Brandon","brandon123ABC"))
tools.assert_false(login(user_accounts, log_in,"BrandonK","123abcABC"))
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: update, points: 0, grade: false, solution: true -->
```python
def update(bank, log_in, username, amount):
    if log_in[username] != True:
        return False

    if username not in bank:
        bank[username] = 0

    if amount + bank[username] < 0:
        return False
    else:
        bank[username] += amount
        return True
```

<!-- grade_id: test_update, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(update(bank,log_in,"Jack",100))
login(user_accounts, log_in, "Brandon", "brandon123ABC")
tools.assert_false(update(bank,log_in,"Brandon",-400))
tools.assert_true(update(bank,log_in,"Brandon",100))
tools.assert_almost_equal(bank.get("Brandon"),215.5)

signup(user_accounts, log_in, "BrandonK", "123aABCD")
tools.assert_is_none(bank.get("BrandonK"))
login(user_accounts,log_in,"BrandonK","123aABCD")
tools.assert_true(update(bank,log_in,"BrandonK",100))
tools.assert_almost_equal(bank.get("BrandonK"),100)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: transfer, points: 0, grade: false, solution: true -->
```python
def transfer(bank, log_in, userA, userB, amount):
    if userA not in bank or log_in[userA] != True:
        return False
    if userB not in log_in:
        return False
    if userB not in bank:
        bank[userB] = 0
    if bank[userA] - amount < 0:
        return False
    else:
        bank[userA] -= amount
        bank[userB] += amount
        return True

```

<!-- grade_id: test_transfer, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(transfer(bank,log_in,"BrandonK","Jack",100))
tools.assert_false(transfer(bank,log_in,"Brandon","JackC",100))
tools.assert_false(transfer(bank,log_in,"Brandon","Jack",100))

login(user_accounts,log_in,"Brandon","brandon123ABC")
tools.assert_false(transfer(bank,log_in,"Brandon","Jack",200))
tools.assert_true(transfer(bank,log_in,"Brandon","Jack",10))
tools.assert_almost_equal(bank.get("Brandon"),105.5)
tools.assert_almost_equal(bank.get("Jack"),55)

signup(user_accounts,log_in,"BrandonK","123aABCD")
tools.assert_is_none(bank.get("BrandonK"))
login(user_accounts,log_in,"BrandonK","123aABCD")
tools.assert_true(transfer(bank,log_in,"Brandon","BrandonK",10))
tools.assert_almost_equal(bank.get("Brandon"),95.5)
tools.assert_almost_equal(bank.get("BrandonK"),10)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: change_password, points: 0, grade: false, solution: true -->
```python
def change_password(user_accounts, log_in, username, old_password, new_password):
    if username not in user_accounts:
        return False
    elif log_in[username] != True:
        return False
    elif user_accounts[username] != old_password:
        return False
    elif old_password == new_password:
        return False
    elif len(new_password) >= 8 and check_upper(new_password) and check_lower(new_password) and check_number(new_password) and username != new_password:
        user_accounts[username] = new_password
        return True
```

<!-- grade_id: test_change_password, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(change_password(user_accounts,log_in,"BrandonK","123abcABC","123abcABCD"))
tools.assert_false(change_password(user_accounts,log_in,"Brandon","brandon123ABC","123abcABCD"))

login(user_accounts,log_in,"Brandon","brandon123ABC")
tools.assert_false(change_password(user_accounts,log_in,"Brandon","123abcABCD","123abcABCDE")) 

tools.assert_false(change_password(user_accounts,log_in,"Brandon","brandon123ABC","brandon123ABC"))

tools.assert_false(change_password(user_accounts,log_in,"Brandon","brandon123ABC","123ABCD"))   

tools.assert_true(change_password(user_accounts,log_in,"Brandon","brandon123ABC","123abcABCD"))
tools.assert_equal("123abcABCD",user_accounts["Brandon"])
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: delete_account, points: 0, grade: false, solution: true -->
```python
def delete_account(user_accounts, log_in, bank, username, password):
    if username not in user_accounts or password != user_accounts[username] or log_in[username] != True:
        return False
    else:
        del user_accounts[username]
        del log_in[username]
        del bank[username]
        return True
```

<!-- grade_id: test_delete_account, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

tools.assert_false(delete_account(user_accounts,log_in,bank,"BrandonK","123abcABC"))
tools.assert_false(delete_account(user_accounts,log_in,bank,"Brandon","123abcABDC"))
tools.assert_false(delete_account(user_accounts,log_in,bank,"Brandon","brandon123ABC"))
login(user_accounts,log_in,"Brandon","brandon123ABC")

tools.assert_true(delete_account(user_accounts,log_in,bank,"Brandon","brandon123ABC"))
tools.assert_is_none(user_accounts.get("Brandon"))
tools.assert_is_none(log_in.get("Brandon"))
tools.assert_is_none(bank.get("Brandon"))
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: main, points: 0, grade: true, solution: false -->
```python
def main():
    '''
    The main function is a skeleton for you to test if your overall programming is working.
    Note we will not test your main function. It is only for you to run and interact with your program.
    '''

    bank = import_and_create_bank("bank.txt")
    user_accounts, log_in = import_and_create_accounts("user.txt")

    while True:
        # for debugging
        print('bank:', bank)
        print('user_accounts:', user_accounts)
        print('log_in:', log_in)
        print('')
        #

        option = input("What do you want to do?  Please enter a numerical option below.\n"
        "1. login\n"
        "2. signup\n"
        "3. change password\n"
        "4. delete account\n"
        "5. update amount\n"
        "6. make a transfer\n"
        "7. exit\n")
        if option == "1":
            username = input("Please input the username\n")
            password = input("Please input the password\n")

            # add code to login
            login(user_accounts, log_in, username, password);
        elif option == "2":
            username = input("Please input the username\n")
            password = input("Please input the password\n")

            # add code to signup
            signup(user_accounts, log_in, username, password)
        elif option == "3":
            username = input("Please input the username\n")
            old_password = input("Please input the old password\n")
            new_password = input("Please input the new password\n")

            # add code to change password
            change_password(user_accounts, log_in, username, old_password, new_password)
        elif option == "4":
            username = input("Please input the username\n")
            password = input("Please input the password\n")

            # add code to delete account
            delete_account(user_accounts, log_in, bank, username, password)
        elif option == "5":
            username = input("Please input the username\n")
            amount = input("Please input the amount\n")
            try:
                amount = float(amount)

                # add code to update amount
                update(bank, log_in, username, amount)
            except:
                print("The amount is invalid. Please reenter the option\n")

        elif option == "6":
            userA = input("Please input the user who will be deducted\n")
            userB = input("Please input the user who will be added\n")
            amount = input("Please input the amount\n")
            try:
                amount = float(amount)

                # add code to transfer amount
                transfer(bank, log_in, userA, userB, amount)
            except:
                print("The amount is invalid. Please re-enter the option.\n")
        elif option == "7":
            break;
        else:
            print("The option is not valid. Please re-enter the option.\n")

#This will automatically run the main function in your program
#Don't change this
if __name__ == '__main__':
    main()
```

Output:
```
bank: {'Brandon': 115.5, 'Patrick': 18.9, 'Sarah': 827.43, 'Jack': 45.0, 'James': 128.87}
user_accounts: {'Brandon': 'brandon123ABC', 'Jack': 'jack123POU', 'James': '100jamesABD', 'Sarah': 'sd896ssfJJH'}
log_in: {'Brandon': False, 'Jack': False, 'James': False, 'Sarah': False}

What do you want to do?  Please enter a numerical option below.
1. login
2. signup
3. change password
4. delete account
5. update amount
6. make a transfer
7. exit
1
Please input the username
Brandon
Please input the password
115.5
bank: {'Brandon': 115.5, 'Patrick': 18.9, 'Sarah': 827.43, 'Jack': 45.0, 'James': 128.87}
user_accounts: {'Brandon': 'brandon123ABC', 'Jack': 'jack123POU', 'James': '100jamesABD', 'Sarah': 'sd896ssfJJH'}
log_in: {'Brandon': False, 'Jack': False, 'James': False, 'Sarah': False}

What do you want to do?  Please enter a numerical option below.
1. login
2. signup
3. change password
4. delete account
5. update amount
6. make a transfer
7. exit
7
```

```python

```